```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

# Homework 5A: Sampling
## Due Date: Monday, July 10th, 11:59 PM

You must submit this assignment to Gradescope by the on-time deadline, Monday, July 10th, 11:59 PM. Please read the syllabus for the grace period policy. No late
submissions beyond the grace period will be accepted. While course staff is happy to help you if you encounter difficulties with submission, we may not be able to respond to last-minute requests for assistance (TAs need to sleep, after all!). **We strongly encourage you to plan to submit your work to Gradescope several hours before the stated deadline.** This way, you will have ample time to reach out to staff for submission support. 

This is part of a two-part assignment. After completing this part ("Homework 5A"), please read all instructions carefully to combine manually graded questions from Homework 5A with Homework 5B, and submit your work to both the coding and written portals of Gradescope.

## Collaboration Policy

Data science is a collaborative activity. While you may talk with others about
the homework, we ask that you **write your solutions individually**. If you do
discuss the assignments with others please **include their names** at the top
of your notebook.

**Collaborators**: *list collaborators here*

## This Assignment

The purpose of this assignment is for you to understand the process of sampling: including convenience samples, random samples; as well as sampling biases. We will also practice drawing samples from a categorical distribution using Numpy.

## Score Breakdown
Question | Manual? | Points
--- | --- | --
1ai | No | 1
1aii | No | 1
1bi | No | 1
1bii | No | 1
1biii | No | 1
1biv | No | 1
2a | No | 2
2b | Yes | 2
2c | Yes | 3
2d | Yes | 2
Total | 3 | 15


### Initialize your environment

```python
import numpy as np
np.random.seed(200)
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
plt.style.use('fivethirtyeight')
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Data 100 Cutest Dog Contest

Welcome to the Data 100 Cutest Dog Contest, Summer 2023 edition! Course staffs nominate their dogs to participate in this contest. Students will vote on the cutest per dog among the nominations in the final exam.

The nominees are:

<img src='images/Shadow.JPG' width="200px" /> <img src='images/Toto.jpeg' width="180px" /> <img src='images/Beck.jpg' width="180px" />

(From left to right) **Shadow** (Alan's sister's dog) , **Toto** (Celine's dog) , and **Beck** (Mihran's dog).

**Earlier in the semester, before the actual contest**, however, course staff would like to predict the results by surveying students in the class. This process is similar to polling that occurs before a political election.

In this section, you are going to explore different sampling methods.
<br/><br/>


<hr style="border: 1px solid #fdb515;" />

## Question 1

<br><br>

---

### Question 1a

Since his dog, Beck, is nominated, Mihran would like to understand the class opinion before the contest takes place. In this coming week, he decides to survey all students enrolled in Data 100 this summer by sending out an email via bCourses asking them to choose the cutest from the three dogs. You may assume no other students/users receive the survey. Mihran closes the survey 12 hours after sending it out.

You can assume that all, and only, enrolled students are on bCourses.

**Part 1**. In Mihran's survey, which of the following is the population of interest? Assign your answer choice to `q1ai` (e.g., `'A'`, `'B'`, etc.).

A. UC Berkeley students<br/>
B. All students enrolled in Data 100</br>
C. All students enrolled in Mihran's discussion section<br/>
D. All students who fill out Mihran's survey

```python
q1ai = ...
```

```python
grader.check("q1ai")
```


**Part 2**. In Mihran's survey, which of the following is the sampling frame? Assign your answer choice to `q1aii` (e.g., `'A'`, `'B'`, etc.).

A. UC Berkeley students<br/>
B. All students enrolled in Data 100</br>
C. All students enrolled in Mihran's discussion section<br/>
D. All students who fill out Mihran's survey

```python
q1aii = ...
```

```python
grader.check("q1aii")
```



**Part 3**. Which of the following is the sample? Assign your answer choice to `q1aiii` (e.g., `'A'`, `'B'`, etc.).

A. UC Berkeley students<br/>
B. All students enrolled in Data 100</br>
C. All students enrolled in Mihran's discussion section<br/>
D. All students who fill out Mihran's survey

```python
q1aiii = ...
```

```python
grader.check("q1aiii")
```

<br><br>

---

### Question 1b

In practice, we cannot get 100\% survey response rate, often because our population is too large, or because there is a time limit. In this case, very few students answered Mihran's survey before he closes it.

To get more data to predict the answer to the original question ("Which dog will win the Data 100 Cutest Dog Contest?"), Mihran decides on a different strategy: **he conducts the pre-contest survey in-person in his discussion section that same week**. He then asks every student who attends discussion that week for their opinion on the cutest of the three pets, by presenting the following slide: 

<img src='images/survey.png' width="600px" />



**Part I:** In this sampling scheme, which of the following is the population of interest? Assign your answer choice to `q1bi` (e.g., `'A'`, `'B'`, etc.).

A. UC Berkeley students<br/>
B. All students enrolled in Data 100</br>
C. All students enrolled in Mihran's discussion section<br/>
D. All students who fill out Mihran's survey

```python
q1bi = ...
```

```python
grader.check("q1bi")
```


**Part II:** In this sampling scheme, which of the following is the sampling frame? Assign your answer choice to `q2bii` (e.g., `'A'`, `'B'`, etc.).

A. UC Berkeley students<br/>
B. All students enrolled in Data 100</br>
C. All students enrolled in Mihran's discussion section<br/>
D. All students who fill out Mihran's survey

```python
q1bii = ...
```

```python
grader.check("q1bii")
```


**Part III:** Which of the following is the sample? Assign your answer choice to `q2biii` (e.g., `'A'`, `'B'`, etc.).

A. UC Berkeley students<br/>
B. All students enrolled in Data 100</br>
C. All students enrolled in Mihran's discussion section<br/>
D. All students who fill out Mihran's survey

```python
q1biii = ...
```

```python
grader.check("q1biii")
```

**Part IV:** Which of the following best characterizes the sample? Assign your answer choice to `q1biv`  (e.g., `'A'`, `'B'`, etc.)..

A. Convenience Sample<br/>
B. Simple Random Sample<br/>
C. Probability Sample

```python
q1biv = ...
```

```python
grader.check("q1biv")
```

<br><br>

<hr style="border: 1px solid #fdb515;" />


## Question 2
Mihran was able to sample 50 students in his section; the results are as follows:

Dog | Vote Share
--- | ---
Shadow | 2%
Toto | 20%
Beck | 78%

Based on this result, he predicts that his dog Beck will win the contest.

Fast-forward to the end of the semester, where the contest has actually taken place. Assume that after tallying every student's votes, **the true popularity of each pet dog** is: 

Dog | Vote Share
--- | ---
Shadow | 10%
Toto | 82%
Beck | 8%

The true winner was actually Toto! Mihran was devastated: what went wrong?

---

### Question 2a

Toto ultimately wins the Cutest Pet Dog Contest, but from Mihran's results, it was predicted that Beck would win. Perhaps there was some underlying issue with Mihran's sampling method.

Probability samples can help us quantify sampling bias and chance error. Put briefly, if we assume that a sample distribution was selected at random from a known population, then we can quantify how likely that sample is to have arisen due to random chance (**chance error**). If the difference in sample and population distributions is too great, then we suspect that the given sample has **bias** in how it was selected from the population.

Write one line of code that runs 1000 independent simulations, where each simulation finds the proportion of voters who voted **Toto** for in a sample of size 50, selected uniformly at random from the **true population**. You may assume that the true population is large enough such that the sample is a random sample with replacement. The output `samples` should be an array with 1000 elements, each of which is the proportion of **Toto** votes in that simulated sample.

**Hint**: 

* Use `np.random.multinomial` ([documentation](https://numpy.org/doc/stable/reference/random/generated/numpy.random.multinomial.html)).
* Use `array[:, i]` to select the ith column of a 2D numpy array
* Feel free to print out `samples` or `samples.shape` to check the output!

```python
samples = ...
samples[:5]
```

```python
grader.check("q2a")
```

<!-- BEGIN QUESTION -->

<br><br>

---

### Question 2b

Create a plot using any `seaborn` and/or `matplotlib.pyplot` functions of your choice to visualize `samples`, which is the simulated distribution of a Toto's vote in a sample of size 50. Include descriptive title and labels. An example is included below. The total area under the plot must be normalized to 1. Your plot may not match exactly to ours due to randomness of the data generating process in `np.random.multinomial`.

<img src='images/q2b.png' width="500px" />

```python
...
plt.show()
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br><br>

---

### Question 2c

According to Mihran's 50-person sample, 20\% of his discussion section reported that they would vote for Toto in the end-of-semester contest.

In the cell below, create a plot using any `seaborn` and/or `matplotlib.pyplot` functions of your choice to visualize Mihran's sample statistic superimposed on the simulated sample distribution you plotted in the previous part. In other words, include
- a vertical line that passes through 20\%, 
- a vertical line that passes through the mean fo the simulated sample distribution, and 
- the simulated sample distribution itself. 

You should choose contrasting colors and include descriptive title, labels and a legend if needed. An example is included below.

<img src='images/q2c.png' width="500px" />

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br><br>

---

### Question 2d

Based on your analysis above, could Mihran's result have arisen due to chance alone? If not, what could be a potential source of bias?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 5A!

### Submission Instructions

Below, you will see two cells.
* Running the first cell will automatically generate a PDF of your answers to all questions that need to be manually graded.
* Running the second cell will automatically generate a zip with your autograded answers.

**You are responsible for combining this resulting Homework 5A PDF with your answers to Homework 5B**, then submitting **both** the coding and written portions of Homework 5 to their respective Gradescope portals:
* **Homework 05 Coding**: Submit your Jupyter notebook zip file for Homework 5A,
which can be generated and downloaded from DataHub by using the `grader.export()`
cell provided below.
* **Homework 05 Written**: Submit a single PDF to Gradescope that contains both
(1) your answers to all manually graded questions from this Homework 5A Jupyter
Notebook, and (2) your answers to all questions in Homework 5B.

To receive credit on this assignment, **you must submit both your coding and written
portions to their respective Gradescope portals.** Your written submission (a single
PDF) can be generated as follows:

1. Access your answers to manually graded Homework 5A  questions in one of three ways:
      * <b>*Automatically create PDF (recommended)*</b>: Run the first cell below and download the generated PDF. This function will extract your response to the manually graded questions and put them on separate page. This process may fail if your answer is not properly formatted; if this is the case, check out common errors and solution described on Ed or follow either of the two ways described below.
      * *Manually download PDF*: If there are issues with automatically generating the PDF in the first cell, you can try downloading the notebook as a PDF by clicking on `File -> Save and Export Notebook As... -> PDF`. If you choose to go this route, you must take special care to ensure all appropriate pages are chosen for each question on Gradescope.
      * *Take screenshots*: If that doesn't work either, you can take screenshots of your answers (and your code if present) to manually graded questions and include them as images in a PDF. The manually graded questions are listed at the top of the Homework 1A notebook.
    
1. Answer the Homework 5B written questions.

1. Combine these two sets of answers together into the same PDF, and submit to the appropriate Gradescope written portal. You can use PDF merging tools, e.g., [Adobe Reader](https://get.adobe.com/reader/), [Smallpdf](https://smallpdf.com/merge-pdf) or [Apple Preview](https://support.apple.com/en-us/HT202945).

1. **Important**: When submitting on Gradescope, you **must tag pages to each question correctly** (it prompts you to do this after submitting your work). This significantly streamlines the grading process for our readers. Failure to do this may result in a score of 0 for untagged questions.

**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

```python
from otter.export import export_notebook
from os import path
from IPython.display import display, HTML
export_notebook("hw05.ipynb", filtering=True, pagebreaks=True)
if(path.exists('hw05.pdf')):
    display(HTML("Download your PDF <a href='hw05.pdf' download>here</a>."))
else:
    print("\n Pdf generation fails, please try the other methods described above")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1ai (points: 1)
```
>>> assert q1ai.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q1aii (points: 1)
```
>>> assert q1aii.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q1aiii (points: 1)
```
>>> assert q1aiii.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q1bi (points: 1)
```
>>> assert q1bi.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q1bii (points: 1)
```
>>> assert q1bii.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q1biii (points: 1)
```
>>> assert q1biii.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q1biv (points: 1)
```
>>> assert q1biv.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q2a (points: 2)
```
>>> assert all(samples <= 1)
>>> assert all(samples >= 0)
>>> assert (len(samples) == 1000)
```
